![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and Pedro Angelo Medeiros Fonini
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
```

# Teste 1: Bisseção

## Questão 1: Número de iterações da bisseção

Generalize o algoritmo da bisseção _iterativa_ para retornar,
além da (aproximação da) raiz, o número de bisseções (divisões por 2 do intervalo) feitas.

Além disso, inclua dois critérios de parada: `xtol` e `ytol`.
O algoritmo termina quando:
- seja possível garantir que o erro (absoluto) da resposta ("em $x$") seja menor do que `xtol`; OU
- quando encontra-se um ponto $p$ tal que $y = f(p)$ (o "erro em $y$") seja menor do que ou igual a `ytol`.

<!-- grade_id: cell-9c2f0bf1fe48da91, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, xtol=1e-8, ytol=1e-8):
    niters=0
    while abs((a-b)/2)>=xtol and abs(f((a+b)/2))>ytol:
        m = (a+b)/2
        if f(m)*f(a)<0:
            b = m
        else:
            a = m
        niters+=1
    return (a+b)/2, niters
```

### 1.1: Testando com $\sqrt2$:

```python
def p2(x): return x**2 - 2
```

```python
bissecao(p2,0,2)
```

Output:
```
(1.4142135605216026, 27)
```

<!-- grade_id: cell-6216e0119f2a0e78, points: 1, grade: true, solution: false -->
```python
raiz, niters = bissecao(p2, 0, 2)
assert abs(raiz - np.sqrt(2)) < 1e-8
assert 25 <= niters <= 30
```

```python
bissecao(p2,0,2,xtol=1e-3)
```

Output:
```
(1.4150390625, 10)
```

<!-- grade_id: cell-caee585cfc6b637f, points: 1, grade: true, solution: false -->
```python
raiz, niters = bissecao(p2, 0, 2, xtol=1e-3)
assert abs(raiz - np.sqrt(2)) < 1e-3
assert 9 <= niters <= 11
```

```python
bissecao(p2, 0, 2, ytol=1e-3)
```

Output:
```
(1.4140625, 7)
```

<!-- grade_id: cell-ed17d59cca42278d, points: 1, grade: true, solution: false -->
```python
raiz, niters = bissecao(p2, 0, 2, ytol=1e-3)
assert abs(raiz**2 - 2) < 1e-3
assert 5 <= niters <= 10
```

<!-- grade_id: cell-f5d73d084a9eb5e6, points: 1, grade: true, solution: false -->
```python
raiz, niters = bissecao(p2, 0, 2, ytol=1e-4)
assert abs(raiz**2 - 2) < 1e-4
assert 10 <= niters <= 15
```

<!-- grade_id: cell-19379484bc5a3adf, points: 1, grade: true, solution: false -->
```python

```

### 1.2: Tolerância, número de iterações, respostas, ...

Observe o seguinte código:

```python
for err in np.logspace(-8, 0, num=9):
    r, n = bissecao(p2, 0, 2, xtol=err)
    print(f"xtol = {err:.1e} --> raiz = {r:.10f}, #iter = {n:3d}")
```

Output:
```
xtol = 1.0e-08 --> raiz = 1.4142135605, #iter =  27
xtol = 1.0e-07 --> raiz = 1.4142135978, #iter =  24
xtol = 1.0e-06 --> raiz = 1.4142141342, #iter =  20
xtol = 1.0e-05 --> raiz = 1.4142074585, #iter =  17
xtol = 1.0e-04 --> raiz = 1.4142456055, #iter =  14
xtol = 1.0e-03 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-02 --> raiz = 1.4140625000, #iter =   7
xtol = 1.0e-01 --> raiz = 1.4375000000, #iter =   4
xtol = 1.0e+00 --> raiz = 1.5000000000, #iter =   1
```

Comente

<!-- grade_id: cell-856d87a71e8f8ec5, points: 2, grade: true, solution: true -->
Podemos ver que ao aumentar a tolerancia do erro, ou seja, ao permitir que o erro seja maior, a precisão da raiz diminui e o número de iterações também.

Como seria o equivalente para o erro em $y$?

<!-- grade_id: cell-c6d45f6ca0f539e4, points: 1, grade: true, solution: true -->
```python
for err in np.logspace(-8, 0, num=9):
    r, n = bissecao(p2, 0, 2, ytol=err)
    print(f"xtol = {err:.1e} --> raiz = {r:.10f}, #iter = {n:3d}")
```

Output:
```
xtol = 1.0e-08 --> raiz = 1.4142135605, #iter =  27
xtol = 1.0e-07 --> raiz = 1.4142135382, #iter =  23
xtol = 1.0e-06 --> raiz = 1.4142136574, #iter =  21
xtol = 1.0e-05 --> raiz = 1.4142150879, #iter =  15
xtol = 1.0e-04 --> raiz = 1.4141845703, #iter =  13
xtol = 1.0e-03 --> raiz = 1.4140625000, #iter =   7
xtol = 1.0e-02 --> raiz = 1.4140625000, #iter =   7
xtol = 1.0e-01 --> raiz = 1.4375000000, #iter =   4
xtol = 1.0e+00 --> raiz = 1.0000000000, #iter =   0
```

Comente, também.

<!-- grade_id: cell-b7931432898e9741, points: 2, grade: true, solution: true -->
Podemos ver novamente a mesma tendência porem o número de iterações diminúi ainda mais para esse caso muito provavelmente devido ao comportamento da função.

### 1.3: Agora, um polinômio um pouco mais complicado:

```python
def p4(x): return x**4 - 3*x + 1
```

<!-- grade_id: cell-b8f933e77c525f69, points: 1, grade: true, solution: false -->
```python
r1, n1 = bissecao(p4, 0, 1)
assert p4(r1) < 1e-6
```

<!-- grade_id: cell-fab8ee9c8e55f544, points: 1, grade: true, solution: false -->
```python
r2, n2 = bissecao(p4, 1, 2)
assert n1 == n2
assert p4(r2) < 1e-6
```

<!-- grade_id: cell-faf0c1fe3a384747, points: 1, grade: true, solution: false -->
```python
r3, n3 = bissecao(p4, 1, 3)
assert n3 == n2+1
assert r3 == r2
```

<!-- grade_id: cell-1ec4a98add6ac67a, points: 1, grade: true, solution: false -->
```python

```

### 1.4: Bônus

Como você faria para calcular a soma e o produto das outras duas raízes de $p_4$?

<!-- grade_id: cell-2c00a397608b6786, points: 2, grade: true, solution: true -->
```python

```

## Questão 2: Função inversa

Nesta questão, vamos implementar uma função inversa, o arco seno, pelo método da bisseção.

A ideia é relativamente simples: sabemos que o arco seno de qualquer ângulo está no intervalo $[-\pi/2, \pi/2]$,
e que a função seno é crescente neste intervalo.

Para encontrarmos uma raiz de $\sin(x) = y_0$, com $y_0$ dado,
devemos construir uma função _auxiliar_ `aux(x)` cuja raiz desejamos calcular pelo método da bisseção.
Implemente a função `arcoseno(y)`, incluindo os argumentos de tolerância para passar para o método da bisseção.

<!-- grade_id: cell-98320234f374ba85, points: 0, grade: false, solution: true -->
```python
def arcoseno(y, xtol=1e-8, ytol=1e-8):
    def aux(x): return np.sin(x)-y
    return bissecao(aux,-np.pi/2,np.pi/2,xtol,ytol)
```

```python
arcoseno(1/2)
```

Output:
```
(0.523598767796069, 26)
```

### 2.1: Testes rápidos

<!-- grade_id: cell-44f6df893329ee62, points: 0, grade: true, solution: false -->
```python
r, n = arcoseno(1/2)
assert abs(np.sin(r) - 1/2) < 1e-8
assert 25 <= n <= 30
```

```python
arcoseno(1)
```

Output:
```
(1.5707004529956536, 14)
```

<!-- grade_id: cell-283d03b9cff34bc9, points: 0, grade: true, solution: false -->
```python
r, n = arcoseno(0.999)
assert abs(np.sin(r) - 0.999) < 1e-8
assert 10 <= n <= 20
```

Como explicar a variação do número de iterações acima?

<!-- grade_id: cell-576d1cfb9243dc18, points: 2, grade: true, solution: true -->
Pode ser explicado pois como o meio do intervalo de -pi/2 e pi/2 é 0 e sen(0) é 1 então a bisseção chega muito próximo da resposta muito rápido, comparado a outra.

### Programando uma função inversa universal

O procedimento que fizemos para o seno pode ser generalizado
para inverter qualquer função contínua monótona num intervalo $[a,b]$.

<!-- grade_id: cell-b082ae8bf3c15203, points: 0, grade: false, solution: true -->
```python
def inversa(f, y, a=0, b=1, xtol=1e-8, ytol=1e-8):
    """Retorna uma aproximação da solução  x  de  f(x) = y no intervalo [a,b]."""
    def aux(x): return f(x)-y
    return bissecao(aux,a,b,xtol,ytol)
```

<!-- grade_id: cell-6315167a64a46646, points: 1, grade: true, solution: false -->
```python
r, n = inversa(np.cos, 1/2, a=0, b=np.pi)
assert abs(np.cos(r) - 1/2) < 1e-8
assert 25 <= n <= 30
```

```python
def f5(x): return x**5 + x + 1
```

<!-- grade_id: cell-cf3670493a5b3d8d, points: 1, grade: true, solution: false -->
```python
r, n = inversa(f5, 1/2, a=-1, b=1)
assert abs(f5(r) - 1/2) < 1e-8
assert 25 <= n <= 30
```

<!-- grade_id: cell-86f691970033c169, points: 1, grade: true, solution: false -->
```python
r, n = inversa(f5, 4, a=-1, b=6)
assert abs(f5(r) - 4) < 1e-7
assert 25 <= n <= 30
```

### 2.3: Retornando a função inversa

Agora, escreva a função `inv(f, a, b)` que retorna a função inversa de $f$ no intervalo $[a,b]$.
Você pode supor que (quem chamar `inv` vai garantir que) $f$ seja monótona neste intervalo.

<!-- grade_id: cell-0faf3256b8c34e85, points: 0, grade: false, solution: true -->
```python
def inv(f, a=0, b=1, xtol=1e-8, ytol=1e-8):
    """Retorna a função inversa de  f  no intervalo [a,b].
    A função inversa é garantida apenas para valores de  y  entre  f(a) e f(b)."""
    def y0(y):
        return inversa(f, y, a, b, xtol, ytol)[0]
    return y0
```

```python
acos = inv(np.cos, 3*np.pi, 4*np.pi, xtol=1e-12, ytol=1e-12)
r = acos(0.5)
np.cos(r)
```

Output:
```
0.5000000000008221
```

<!-- grade_id: cell-013c35404c6a69b8, points: 1, grade: true, solution: false -->
```python
acos = inv(np.cos, 3*np.pi, 4*np.pi, xtol=1e-12, ytol=1e-12)
r = acos(0.5)
assert abs(np.cos(r) - 0.5) <= 1e-12
```

<!-- grade_id: cell-99bb51fcb9dde95d, points: 1, grade: true, solution: false -->
```python
r = acos(0)
assert abs(np.cos(r)) <= 1e-12
assert abs(r - 7*np.pi/2) <= 1e-11
```

<!-- grade_id: cell-bf7aaaa0cfed0694, points: 2, grade: true, solution: false -->
```python
inv_f5 = inv(f5, 0, 10, xtol=1e-14, ytol=1e-14)
ys = [2,3,4,5,6]
xs = [inv_f5(y) for y in ys]
for x,y in zip(xs,ys):
    assert np.abs(f5(x) - y) < 1e-13
```

## Questão 3: Bisseção em listas ordenadas

Se desejamos encontrar um número em uma lista, podemos percorrer cada um de seus índices
até achar o número, e retornamos o índice correspondente.

Mas, se a lista `l` possui números em ordem crescente,
é possível usar uma variante da bisseção para encontrar este índice.
Observe que agora desejamos retornar um número **inteiro**,
e que índices para listas sempre devem ser inteiros, portanto:
- cuidado ao escolher o "ponto médio"
- mas o critério de parada é "quando achar o índice".

### 3.1: Encontrando um número

<!-- grade_id: cell-8c9f49b0ba8b5d04, points: 0, grade: false, solution: true -->
```python
def bissect_list(l, v):
    """Índice  k  na lista  l , crescente, tal que l[k] = v.  None caso não exista."""
    n = len(l)
    a = 0
    if len(l) == 0 :
        return None
    if v not in l:
        return None
    while abs(n-a)>=2:
        if (n+a)%2 == 0:
            m = int((n+a)/2)
            if l[m]>v:
                n=m
            elif l[m]<v:
                a=m
            else:
                return m
        else:
            m = int((n+a-1)/2)
            if l[m]>v:
                n=m
            elif l[m]<v:
                a=m
            else:
                return m
    if l[a] == v:
        return a
    else:
        return n
        
```

<!-- grade_id: cell-d0e135df902b3955, points: 1, grade: true, solution: false -->
```python
assert bissect_list([], 42) is None
```

<!-- grade_id: cell-aa0a59b135a65159, points: 1, grade: true, solution: false -->
```python
l = [4]
assert bissect_list(l, 4) == 0
```

<!-- grade_id: cell-ba03109833002639, points: 1, grade: true, solution: false -->
```python
assert bissect_list(l, 5) is None
```

<!-- grade_id: cell-84b5ba3a3384d223, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28]
for i,v in enumerate(l):
    idx = bissect_list(l, v)
    assert idx == i, f"v={v}, i={i}, idx={idx}"
```

### 3.2: Encontrando o primeiro valor maior ou igual ao número dado

Escreva uma nova bisseção, agora para encontrar o primeiro índice `k` tal que
`l[k] >= v`.

<!-- grade_id: cell-e442fe23480d4b8d, points: 0, grade: false, solution: true -->
```python
def bissect_geq(l, v):
    """First index  k  on an increasing list  l  such that  l[k] >= v.
    Returns  len(l)  if  l[-1] < v."""
    n = len(l)-1
    a = 0
    r = 0
    if len(l) == 0 :
        return None
    if  l[-1] < v:
        return len(l)
    if l[a] >= v:
        return a
    if l[n] == v:
        return n
    while abs(n-a)>=2:
        if (n+a)%2 == 0:
            m = int((n+a)/2)
            if l[m]>v:
                n=m
                r=m
            elif l[m]<v:
                a=m
            else:
                return m
        else:
            m1 = int((n+a-1)/2)
            m2 = int((n+a+1)/2)
            if l[m1] == v:
                return m1
            elif l[m2]==v:
                return m2
            elif l[m2]<v:
                a=m2
            else:
                n=m1
                r=m2
    return r
```

<!-- grade_id: cell-8d0e1c185171dbce, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28,35,36]
assert bissect_geq(l, 12) == 4
```

<!-- grade_id: cell-d78ee0a4881484f1, points: 2, grade: true, solution: false -->
```python
for v in range(50):
    idx = bissect_geq(l, v)
    if idx < len(l):
        assert l[idx] >= v, f"idx={idx}, v={v}, l[idx]={l[idx]}"
    if idx > 0:
        assert l[idx-1] < v, f"idx={idx}, v={v}, l[idx-1]={l[idx-1]}"
```

```python

```