![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and Pedro Angelo Medeiros Fonini
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Matrizes não simétricas e autovalores complexos

## Questão 1: Sub-espaços invariantes

Mostre que uma matriz real $M$ com dois autovalores $\lambda$ e $\bar\lambda$ conjugados,
também possui dois autovetores "complexos conjugados", $u$ e $\bar u$.

<!-- grade_id: cell-acdfa31ca03690b8, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

De $u$ e $\bar u$, é possível extrair dois vetores $x$ e $y$, reais.
Diga como.

<!-- grade_id: cell-e84ded2eafbdcdb4, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Seja $r e^{i \theta} = \lambda$ a forma polar do autovalor $\lambda$.
A partir da equação de autovalor de $u$, separe as partes reais e imaginárias,
e dê, apenas em função de $x$, $y$, $r$ e $\theta$,
os valores de $Ax$ e $Ay$.

Deduza que, no sub-espaço real gerado por $x$ e $y$,
a matriz $A$ age como uma **dilatação e rotação**.

<!-- grade_id: cell-fca9f411b94f5659, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

### Bônus: contra-exemplo

Dê uma matriz $3 \times 3$ tal que o "maior autovalor" seja um par complexo conjugado,
**e** cujos vetores correspondentes $x$ e $y$ **não sejam ortogonais**

<!-- grade_id: cell-f6e44c14bfa651ae, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2: Iteração de potência 2D

Usando as equações da questão 1, mostre que $y$ é uma combinação linear de $x$ e $Ax$
(dependendo de $r$ e $\theta$).

<!-- grade_id: cell-a7ffd642d8777c59, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

A iteração de potência não converge para o maior autovalor (em norma),
afinal, $A$ é real e os vetores que estamos usando também são,
e o autovalor é complexo.

Mostre, entretanto, que a iteração de potência para **2 vetores** (mantidos ortogonais)
converge no seguinte sentido:
- o espaço gerado pelos vetores $a_n$, $b_n$ converge para o subespaço real gerado por $x$ e $y$.

<!-- grade_id: cell-cab719411c3b5a8c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Iteração 3D

O [produto vetorial](https://pt.wikipedia.org/wiki/Produto_vetorial)
de dois vetores $x$ e $y$ em $R^3$ aponta na direção ortogonal ao plano gerado por eles.

<!-- grade_id: cell-f46e44c064bf52d6, points: 0, grade: false, solution: true -->
```python
def vector(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-61661a04abe40af1, points: 1, grade: true, solution: false -->
```python
u = [1,0,0]
v = [0,1,0]
assert np.all(vector(u,v) == [0,0,1])
```

<!-- grade_id: cell-fca8d57520339f10, points: 1, grade: true, solution: false -->
```python
u = [1,2,3]
v = [3,2,1]
assert np.all(vector(u,v) == [-4,8,-4])
assert np.all(vector(u,v) == -vector(v,u))
```

<!-- grade_id: cell-b12002bc93c68649, points: 2, grade: true, solution: false -->
```python
np.random.seed(1)
for _ in range(10):
    u = np.random.randn(3)
    v = np.random.randn(3)
    z = vector(u,v)
    assert np.isclose(np.dot(u,z), 0)
    assert np.isclose(np.dot(v,z), 0)
```

Use a convergência do vetor (normalizado) $z = x \times y$
como critério de parada para a iteração de potência.

Não esqueça de normalizar (e ortogonalizar) os vetores $x$ e $y$ a cada iteração ;-)

<!-- grade_id: cell-99ce97c04eaa9a6c, points: 0, grade: false, solution: true -->
```python
def power_iter_vetorial(A, tol=1e-6, maxiter=1000):
    """Iteração de potência para encontrar um subespaço 2D correspondente a autovalores complexos de uma matriz real.
    Retorna as listas de estimativas para
    - o vetor  x \times y  unitário normal ao subespaço.
    """
    n,m = np.shape(A)
    assert n == m, 'A must be square'
    assert n == 3, 'Only works for 3x3 matrices'

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73a3e9984aa42970, points: 2, grade: true, solution: false -->
```python
np.random.seed(5)
M = np.random.randn(3,3)
zs = power_iter_vetorial(M)

e,U = np.linalg.eig(M)

assert -1e-6 < np.dot(U[:,0].real, zs[-1]) < 1e-6
assert -1e-6 < np.dot(U[:,0].imag, zs[-1]) < 1e-6
```

<!-- grade_id: cell-d3d1c7caa7d2598d, points: 1, grade: true, solution: false -->
```python
np.random.seed(5)
M = np.random.randn(3,3)
zs = power_iter_vetorial(M, tol=1e-12)

e,U = np.linalg.eig(M)

assert -1e-12 < np.dot(U[:,0].real, zs[-1]) < 1e-12
assert -1e-12 < np.dot(U[:,0].imag, zs[-1]) < 1e-12
```

Faça um gráfico com a distância entre os vetores $z$ sucessivos.
Com que velocidade o algoritmo converge?

<!-- grade_id: cell-1df16f6d3b2da04a, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1d85a52bf6f8cf0c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE